Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel→Restart) and then run all cells (in the menubar, select Cell→Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Bi 621 – Problem Set 3
## Due before class, Monday, July 9

The goal of this problem set is to convert phred quality scores from their single letter encoding into integers and to calculate the average phred quality score for a single sequence. **Be sure to remove the ```raise NotImplementedError()``` from the code.**

```python
#initializing variables, do not change
id = "@GAAATG_1_1101_4446_2137_1"
seq = "TGCAGGTTGAGTTCTCGCTGTCCCGCCCCCATCTCTTCTCTTCCAGTCTGGCTCTGGAGCAGTTGAGCCCAGCTCAGGTCCGCCGGAGGAGACCG"
phred_score = "FFHHHHHJJJJIJIJJJIJJJJJJIIIJJJEHJJJJJJJIJIDGEHIJJFIGGGHFGHGFFF@EEDE@C??DDDDDDD@CDDDDBBDDDBDBDD@"
```

# Part 1 – Convert_phred
Write a function that takes a single character as it's parameter and converts that parameter into a phred score, and returns the phred score.
>*Remember: Illumina 1.8+ phred scores are offset by a factor of 33 from 0. Review the lecture notes as well as the [Wikipedia FASTQ page](https://en.wikipedia.org/wiki/FASTQ_format) to see how to do the conversion.*

*HINT* – what does the ord() function do? Check the documentation

<!-- grade_id: convert_phred, points: 0, grade: false, solution: true -->
```python
# Creates a method that takes an instance variable.
def convert_phred(letter):
    """Converts a single character into a phred score"""
    
    # Creates a variable that stores a difference 
    # of -33 of the unicode point from the instance variable. 
    x=ord(letter) - 33
    
    # returns a variable with the result.
    return x
  
```

Check to see that your code works.

```python
# Calls a method. 
convert_phred('A')
```

Output:
```
32
```

<!-- grade_id: correct_convert_phred, points: 3, grade: true, solution: false -->
```python
#autograder tests, do not change
"""Check that convert_phred returns the correct value for several different inputs"""
assert convert_phred("I") == 40
assert convert_phred("C") == 34
assert convert_phred("2") == 17
assert convert_phred("@") == 31
assert convert_phred("$") == 3
```

# Part 2 – Convert the entire phred string
Write a loop print each letter in the `phred_score` string and its corresponding score in the following format:
```
0: F - 37
1: F - 37
2: H - 39
3: H - 39
4: H - 39
5: H - 39
6: H - 39
7: J - 41
...```
Reminder, start simply.
* Write a loop to iterate over each letter in the `phred_score` and print them out.
* Modify your loop to print the letter and the corresponding score. *Be sure to use your `convert_phred` function*.
* Alter the loop to print out an index: 0 for the first letter, 1 for the 2nd letter, etc.
* Be sure your formatting matches the requested output.

<!-- grade_id: print_phred_score, points: 3, grade: true, solution: true -->
```python
# Creates a variable
count = 0

# A for loop that goes through a string.
for letter in phred_score:
    
    # prints a statement in a specific format
    print("{0}: {1} - {2}".format(count, letter, convert_phred(letter) ))
    
    # Adds a count of 1 to the variable.
    count= count + 1

```

Output:
```
0: F - 37
1: F - 37
2: H - 39
3: H - 39
4: H - 39
5: H - 39
6: H - 39
7: J - 41
8: J - 41
9: J - 41
10: J - 41
11: I - 40
12: J - 41
13: I - 40
14: J - 41
15: J - 41
16: J - 41
17: I - 40
18: J - 41
19: J - 41
20: J - 41
21: J - 41
22: J - 41
23: J - 41
24: I - 40
25: I - 40
26: I - 40
27: J - 41
28: J - 41
29: J - 41
30: E - 36
31: H - 39
32: J - 41
33: J - 41
34: J - 41
35: J - 41
36: J - 41
37: J - 41
38: J - 41
39: I - 40
40: J - 41
41: I - 40
42: D - 35
43: G - 38
44: E - 36
45: H - 39
46: I - 40
47: J - 41
48: J - 41
49: F - 37
50: I - 40
51: G - 38
52: G - 38
53: G - 38
54: H - 39
55: F - 37
56: G - 38
57: H - 39
58: G - 38
59: F - 37
60: F - 37
61: F - 37
62: @ - 31
63: E - 36
64: E - 36
65: D - 35
66: E - 36
67: @ - 31
68: C - 34
69: ? - 30
70: ? - 30
71: D - 35
72: D - 35
73: D - 35
74: D - 35
75: D - 35
76: D - 35
77: D - 35
78: @ - 31
79: C - 34
80: D - 35
81: D - 35
82: D - 35
83: D - 35
84: B - 33
85: B - 33
86: D - 35
87: D - 35
88: D - 35
89: B - 33
90: D - 35
91: B - 3
```

Write a second function that takes the original, unmodified `phred_score` string as a parameter. This function should calculate the average quality score of the whole phred string.
> Struggling? See some [hints](PS3_hints.ipynb)

<!-- grade_id: qual_score, points: 0, grade: false, solution: true -->
```python
# Creates a method that takes an instance variable.
def qual_score(phred_score):
    """Calculates a quality score from a phred score"""
    
    # Creates a variable
    count = 0
    total = 0
    
    # A for loop that goes through a string.
    for letter in phred_score:
        
        # Creates a variable to store the total sum of the phred scores.
        total= convert_phred(letter) + total
        
        # Adds a count of 1 to the variable.
        count = count + 1
        
    # returns the average result for the quuality score.
    return (total/count)
```

```python
# Prints a statement.
print(qual_score(phred_score))
```

Output:
```
37.62105263157895
```

<!-- grade_id: qual_score_call_fxn, points: 1, grade: true, solution: false -->
```python
#autograder tests, do not change
old_convert_phred = convert_phred
del convert_phred

try:
    qual_score(phred_score)
except NameError:
    pass
else:
    raise AssertionError("qual_score does not call convert_phred")
finally:
    convert_phred = old_convert_phred
    del old_convert_phred
```

<!-- grade_id: correct_qual_score, points: 3, grade: true, solution: false -->
```python
#autograder tests, do not change
assert qual_score(phred_score) == 37.62105263157895
```

```python

```

```python

```