Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Gera Perez"
COLLABORATORS = ""
```

# ICA7 - Transitioning to Python: Functions

## Write a function to calculate the Pythagorean theorem
$$a^2 + b^2 = c^2$$
1. The function should take 2 parameters (a and b), compute the value of c and return it
2. Call your function with several different values for a and b and print the results
3. Check your work with a calculator
4. After writing your function in Jupyter Notebook, convert it to a .py file and execute it from the command line

Turn in your completed Jupyter Notebook and your .py file to GitHub.

<!-- grade_id: cell-f0ce8715e4e9a546, points: 0, grade: false, solution: true -->
```python
# Imports a module.
import math

# Creates a method that takes two parameters.
def pythag(a,b):
    '''Computes the length of the hypotenuse of a right triangle given the lengths
    of the other two sides.'''
    
    # creates a variable that squares each parameter then computes an addition. 
    x = (a**2) + (b**2) 
    
    # creates a variable that stores the square root of a previous result.
    c = math.sqrt(x)
    
    # Returns a variable that has a result.
    return c
```

```python
# calls a method
pythag(1,2)
```

Output:
```
2.23606797749979
```

<!-- grade_id: cell-90ad85a951d5e183, points: 5, grade: true, solution: false -->
```python
'''Check function results, do not change this cell'''
assert pythag(3,4) == 5, "fails 3,5 test"
assert pythag(5,12) == 13, "fails 5,12 test"
assert pythag(9,40) == 41, "fails 9,40 test"
```

Convert your Jupyter notebook to a script and be sure to upload the script to your repo. Include the name of your script file in the text box below.

<!-- grade_id: cell-84c60331beea15a6, points: 5, grade: true, solution: true -->
ICA7.py

```python

```

```python

```